[1613560682] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613560742] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613560795] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1613560802] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.000 second response time
[1613560832] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613560855] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1613560862] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613560892] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613560915] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1613560952] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613560975] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1613561012] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613566853] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;SOFT;1;connect to address 10.250.10.2 and port 80: Connection refused
[1613566913] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;SOFT;2;connect to address 10.250.10.2 and port 80: Connection refused
[1613566973] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;SOFT;3;connect to address 10.250.10.2 and port 80: Connection refused
[1613567013] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567033] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;HARD;4;connect to address 10.250.10.2 and port 80: Connection refused
[1613567073] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567133] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567193] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613573520] SERVICE ALERT: viewadmin110;Jenkins;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.003 second response time
[1613573580] SERVICE ALERT: viewadmin110;Jenkins;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.002 second response time
[1613573640] SERVICE ALERT: viewadmin110;Jenkins;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.002 second response time
[1613573700] SERVICE ALERT: viewadmin110;Jenkins;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.002 second response time